```python
import numpy as np
import matplotlib.pyplot as plt
```

# Enunciado

## Introducción

A menudo es necesario transmitir varias señales simultáneamente por un único canal. Para ello las señales deben *multiplexarse* . Existen varios tipos de multiplexación, pero los más usuales son:

* Multiplexación en el tiempo: Donde las muestras de las señales se van alternando en el tiempo.

* Multiplexación en la frecuencia: donde cada una de las señales se modula con una portadora de frecuencia diferente.

## Multiplexación en el tiempo

En este ejercicio se realizará un multiplexor de señales en el tiempo:

* Cada una de las señales a multiplexar será un array de numpy de tamaño *(N,)* . Supondremos que todas las señales son de la misma duración.

* Las señales a multiplexar estarán almacenadas en una lista. Supondremos que tenemos *S* señales.

* La función devolverá un array de numpy con la señal multiplex de tamaño *(NxS,1)*

La siguiente celda ilustra el funcionamiento esperado:

```python
x1=np.array([1,2,3,4,5])
x2=np.array([11,12,13,14,15])
x3=np.array([21,22,23,24,25])
senyales=[x1,x2,x3]
print("Senyales a multiplexar:",senyales)
resultado_esperado=np.array([1,11,21, 2,12,22, 3,13,23, 4,14,24, 5,15,25])
print("Resultado esperado de la multiplexación:", resultado_esperado)
```

Output:
```
Senyales a multiplexar: [array([1, 2, 3, 4, 5]), array([11, 12, 13, 14, 15]), array([21, 22, 23, 24, 25])]
Resultado esperado de la multiplexación: [ 1 11 21  2 12 22  3 13 23  4 14 24  5 15 25]
```

<!-- grade_id: cell-ed1f77e16e374b6f, points: 0, grade: false, solution: true -->
```python
def multiplex_temporal(lista_senyales):
    ''' Multiplexación temporal de senyales'''
    ### BEGIN SOLUTION
    N=len(lista_senyales[0])
    S=len(lista_senyales)
    y=np.zeros(N*S)
    for k in range(S):
        y[k::S]=lista_senyales[k]
    return y
    ### END SOLUTION
    
    
```

```python
""" Verificación pre-envío de funcionamiento correcto
Pruebe a ejecutar esta celda hasta que no haya errores
"""
assert multiplex_temporal(senyales) is not None, "Probablemente falta return"
assert type(multiplex_temporal(senyales)) == type(x1), "Debe devolver un array de numpy"
assert resultado_esperado.shape== multiplex_temporal(senyales).shape, "El tamaño de la salida no es válido"
assert np.all(multiplex_temporal(senyales) == resultado_esperado), "El resultado tiene el tamaño adecuado pero no los valores adecuados"

```

<!-- grade_id: cell-d5089a0373aa90d6, points: 2, grade: true, solution: false -->
```python
""" Verificación de funcionamiento correcto"""

### BEGIN HIDDEN TESTS
SS=4
lista=[]
N=100
for k in range(SS):
    lista.append(np.random.rand(N))
resultado=multiplex_temporal(lista)

assert resultado is not None, "Probablemente falta return"
assert type(resultado) == type(lista[0]), "Debe devolver un array de numpy"
assert resultado.shape== (N*SS,), "El tamaño de la salida no es válido"
for k in range(len(resultado)):
    s=k%SS
    n=k//SS
    assert resultado[k]==lista[s][n], "El resultado tiene el tamaño adecuado pero no los valores adecuados"

### END HIDDEN TESTS
```

## Demultiplexación

Cuando las señales se transmiten multiplexadas, es normal tener que volver a separarlas en el receptor. A dicho proceso se le denomina demultiplexión.

En este apartado, realizará una función para realizar la demultiplexión temporal:

* Recibirá un array de numpy de tamaño (NxS,)

* Recibirá el número de canales S

* Devolverá una lista con L elementos cada uno de los cuales será un array de numpy de tamaño N

La siguiente celda ilustra lo que se pretende: 

```python
senyal_multiplexada=np.array([1,11,21, 2,12,22, 3,13,23, 4,14,24, 5,15,25])
print("Senyal multiplexada:", resultado_esperado)
resultado_esperado_demultiplex=[np.array([1,2,3,4,5]), 
                                np.array([11,12,13,14,15]),
                                np.array([21,22,23,24,25])]

print("Resultado esperado demultiplexación :",resultado_esperado_demultiplex)

```

Output:
```
Senyal multiplexada: [ 1 11 21  2 12 22  3 13 23  4 14 24  5 15 25]
Resultado esperado demultiplexación : [array([1, 2, 3, 4, 5]), array([11, 12, 13, 14, 15]), array([21, 22, 23, 24, 25])]
```

Por supuesto, la función realizada **deberá funcionar para cualquier cantidad de canales y duración de señales**

<!-- grade_id: cell-d0e3e704690478d8, points: 0, grade: false, solution: true -->
```python
def demultiplex_temporal(multiplex, S):
    ''' Demultiplexación temporal de senyales
    
    multiplex: muestras de la señal multiplexada
    S: número de canales

    Devuelve: lista de señales correspondientes a los canales.
    '''

    # Primero comprobamos que el tamaño de multiplex es múltiplo de S

    assert multiplex.shape[0]%S == 0, "La señal de entrada debe tener un tamaño múltiplo de S"
    
    ### BEGIN SOLUTION
    out=[]
    for k in range(S):
        out.append(multiplex[k::S])
    return out
    ### END SOLUTION
```

```python
""" Verificación pre-envío:
Comprobar que las siguientes comprobaciones no dan errores"""
multiplex=np.array([1,11,21, 2,12,22, 3,13,23, 4,14,24, 5,15,25])
S=3
resultado=demultiplex_temporal(multiplex,S)
assert resultado is not None, "Probablemente falta return"
assert type(resultado) == list, "Debe devolver una lista"
assert len(resultado)==S, "El resultado debe tener tantos elementos como canales"
for k in range(S):
    assert len(resultado[k])==len(multiplex)//S, "Las longitudes de las señales no son correctas. Deben coincidir con longitud de la señal múltiplex dividida por S"
x1=np.array([1,2,3,4,5])
x2=np.array([11,12,13,14,15])
x3=np.array([21,22,23,24,25])
resultado_esperado=[x1,x2,x3]
for k in range(S):
    assert np.all(resultado[k] == resultado_esperado[k]), "El resultado tiene el tamaño adecuado pero no los valores adecuados"


```

<!-- grade_id: cell-0d3dc2e472854d71, points: 2, grade: true, solution: false -->
```python
""" Verificación de funcionamiento correcto"""

### BEGIN HIDDEN TESTS
SS=7
N=100
multiplex1=np.random.randn(SS*N)
resultado1=demultiplex_temporal(multiplex1,SS)



assert resultado1 is not None, "Probablemente falta return"
assert type(resultado1) == list, "Debe devolver una lista"
assert len(resultado1)== SS, "El tamaño de la lista no es válido"
for k in range(S):
    assert len(resultado1[k])==len(multiplex1)//SS, "Las longitudes de las señales no son correctas. Deben coincidir con longitud de la señal múltiplex dividida por S"
for k in range(len(multiplex1)):
    s=k%SS
    n=k//SS
    assert multiplex1[k]==resultado1[s][n], "El resultado tiene el tamaño adecuado pero no los valores adecuados"

### END HIDDEN TESTS
```